## Tips
- To avoid unpleasant surprises, I suggest you _run all cells in their order of appearance_ (__Cell__ $\rightarrow$ __Run All__).


- If the changes you've made to your solution don't seem to be showing up, try running __Kernel__ $\rightarrow$ __Restart & Run All__ from the menu.


- Before submitting your assignment, make sure everything runs as expected. First, restart the kernel (from the menu, select __Kernel__ $\rightarrow$ __Restart__) and then **run all cells** (from the menu, select __Cell__ $\rightarrow$ __Run All__).

## Reminder

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, UA email, and collaborators below:



Several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
NAME = ""
# University of Arizona email address
EMAIL = ""
# Names of any collaborators.  Write N/A if none.
COLLABORATORS = ""
```

## Scratchpad

You are welcome to create new cells (see the __Cell__ menu) to experiment and debug your solution.

```python
%load_ext autoreload
%autoreload 2
```

# Mini Python tutorial

This course uses Python 3.8.

Below is a very basic (and incomplete) overview of the Python language... 

For those completely new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.8/library/stdtypes.html#common-sequence-operations).

```python
# This is a comment.  
# Any line starting with # will be interpreted as a comment

# this is a string assigned to a variable
greeting = "hello"

# If enclosed in triple quotes, strings can also be multiline:

"""
I'm a multiline
string.
"""

# let's use a for loop to print it letter by letter
for letter in greeting:
    print(letter)
    
# Did you notice the indentation there?  Whitespace matters in Python!

# here's a list of integers

numbers = [1, 2, 3, 4]

# let's add one to each number using a list comprehension
# and assign the result to a variable called res
# list comprehensions are used widely in Python (they're very Pythonic!)

res = [num + 1 for num in numbers]

# let's confirm that it worked
print(res)

# now let's try spicing things up using a conditional to filter out all values greater than or equal to 3...
print([num for num in res if not num >= 3])

# Python 3.7 introduced "f-strings" as a convenient way of formatting strings using templates
# For example ...
name = "Josuke"

print(f"{greeting}, {name}!")

# f-strings are f-ing convenient!


# let's look at defining functions in Python..

def greet(name):
    print(f"Howdy, {name}!")

# here's how we call it...

greet("partner")

# let's add a description of the function...

def greet(name):
    """
    Prints a greeting given some name.
    
    :param name: the name to be addressed in the greeting
    :type name: str
    
    """
    print(f"Howdy, {name}!")
    
# I encourage you to use docstrings!

# Python introduced support for optional type hints in v3.5.
# You can read more aobut this feature here: https://docs.python.org/3.8/library/typing.html
# let's give it a try...
def add_six(num: int) -> int:
    return num + 6

# this should print 13
print(add_six(7))

# Python also has "anonymous functions" (also known as "lambda" functions)
# take a look at the following code:

greet_alt = lambda name: print(f"Hi, {name}!")

greet_alt("Fred")

# lambda functions are often passed to other functions
# For example, they can be used to specify how a sequence should be sorted
# let's sort a list of pairs by their second element
pairs = [("bounce", 32), ("bighorn", 12), ("radical", 4), ("analysis", 7)]
# -1 is last thing in some sequence, -2 is the second to last thing in some seq, etc.
print(sorted(pairs, key=lambda pair: pair[-1]))

# we can sort it by the first element instead
# NOTE: python indexing is zero-based
print(sorted(pairs, key=lambda pair: pair[0]))

# You can learn more about other core data types and their methods here: 
# https://docs.python.org/3.8/library/stdtypes.html

# Because of its extensive standard library, Python is often described as coming with "batteries included".  
# Take a look at these "batteries": https://docs.python.org/3.8/library/

# You now know enough to complete this homework assignment (or at least where to look)
```

```python

```

```python
from typing import Iterator, Iterable, List, Tuple, Text, Union
from sklearn.metrics import f1_score, accuracy_score
import numpy as np

from scipy.sparse import spmatrix
```

```python
# An NDArray can either be a numpy array (np.ndarray) or a sparse matrix (spmatrix)
NDArray = Union[np.ndarray, spmatrix]

np.random.seed(42)
```

# Assignment overview

In this assignment, you'll use a binomial logistic regression classifier with $n$-gram features to categorize SMS messages as **SPAM** or **NOT SPAM**.  

You will not be required to implement logistic regression from scratch.  Instead, you may use the `scikit-learn` implementation.  Be sure to carefully read the comments in the code skeleton provided.  To earn full credit, all tests must pass.

<!-- grade_id: needed-imports, points: 0, grade: false, solution: true -->
```python
# Add any needed imports from sklearn here! See the comments/hints below for
#  tips that will guide you to what to add here.
# YOUR CODE HERE
raise NotImplementedError()
```

## `.read_smsspam()`

First we'll need data for training and evaluating our classifier.

Complete the function below which takes a path to a file (stored under `data` in the docker image) and returns a sequence of (label, text) pairs.

<!-- grade_id: code-read-smsspam, points: 0, grade: false, solution: true -->
```python
def read_smsspam(smsspam_path: Text) -> Iterator[Tuple[Text, Text]]:
    """
    Generates (label, text) tuples from the lines in an SMSSpam file.

    SMSSpam files contain one message per line. Each line is composed of a label
    (ham or spam), a tab character, and the text of the SMS. Here are some
    examples:

      spam	85233 FREE>Ringtone!Reply REAL
      ham	I can take you at like noon
      ham	Where is it. Is there any opening for mca.

    :param smsspam_path: The path of an SMSSpam file, formatted as above.
    :return: An iterator over (label, text) tuples.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Test `.read_smsspam()` (3 pts)

<!-- grade_id: test-read-smsspam, points: 3, grade: true, solution: false -->
```python
# keep a counter here (instead of enumerate) in case the iterator is empty
count = 0
for example in read_smsspam("data/smsspam/SMSSpamCollection.train"):

    # make sure the right shape is returned
    assert len(example) == 2
    label, text = example

    # make sure the label is one of the expected two
    assert label in { "ham", "spam" }

    count += 1

# You should find exactly 3345 pairs in the training partition of the data
assert count == 3345
```

## `TextToFeatures`

Using the provided skeleton, complete the class `TextToFeatures` to transform text (documents) into real-valued features derived from the training data.

<!-- grade_id: code-texttofeatures, points: 0, grade: false, solution: true -->
```python
class TextToFeatures:
    def __init__(self):
        """
        Initializes an object for converting texts to features.    
        """
        
        # HINT: you may want to use a sklearn vectorizer. Be sure you've
        #  worked through the sklearn tutorial in the course website and read
        #  the documentation on sklearn's vectorizers at
        # https://scikit-learn.org/stable/modules/classes.html#module-sklearn.feature_extraction.text
        #
        # Learning to use a library based only on reading its documentation
        #  is one of the skills that we hope you'll learn in this course!
        #
        # If you take this approach, your code here will largely just be
        #  calls to the corresponding methods of the sklearn vectorizer
        #  that you choose. Don't forget to add the import statement
        #  for your chosen vectorizer in the imports cell in the Assignment
        #  Overview.
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def fit(self, training_texts: Iterable[Text]) -> None:
        """
        Fits ("trains") a TextToFeature instance on a collection of documents.
        
        The provided training texts are analyzed to determine the vocabulary, 
        i.e., all feature values that the converter will support. 
        Each such feature value will be associated with a unique integer index 
        that may later be accessed via the .index() method.

        It is up to the implementer exactly what features to produce from a
        text, but the features will always include some single words and some
        multi-word expressions (e.g., "need" and "to you").
        
        
        docs = [
            "LOL. is this u? http://supersketchyurl.com/dangerous",
            "The IRS has been trying to reach you.",
            "Enclosed is your Coyote Joe's Marketplace Rewards Card."
            "Logan I'd like to add you to my professional network on LinkedIn",
        ]
        
        t2f = TextToFeatures()
        t2f.fit(docs)

        :param training_texts: The training texts.
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        
        
    def index(self, feature: Text) -> Union[None, int]:
        """
        Returns the index in the vocabulary of the given feature value.  
        If the features isn't present, return None.

        :param feature: A feature
        :return: The unique integer index associated with the feature or None if not present.
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def transform(self, texts: Iterable[Text]) -> NDArray:
        """
        Creates a feature matrix from a sequence of texts.
        
        docs = [
            "LOL. is this u? http://supersketchyurl.com/dangerous",
            "The IRS has been trying to reach you.",
            "Enclosed is your Coyote Joe's Marketplace Rewards Card."
            "I'd like to add you to my professional network on LinkedIn",
        ]
        
        t2f = TextToFeatures()
        t2f.fit(docs)

        # this produces a NDArray representing our features for the provided doc
        t2f.transform(["Let's meet at Coyote Joe's at 6."])


        Each row of the matrix corresponds to one of the input texts. The value
        at index j of row i is the value in the ith text of the feature
        associated with the unique integer j.

        It is up to the implementer what the value of a feature that is present
        in a text should be, though a common choice is 1. Features that are
        absent from a text will have the value 0.

        :param texts: A sequence of texts.
        :return: A matrix, with one row of feature values for each text.
        """
        # YOUR CODE HERE
        raise NotImplementedError()

```

## Test features (7 pts)

Let's test the behavior of your implementation of `TextToFeatures` (2 points for the first cell, 5 points for the second cell)

<!-- grade_id: test-texttofeatures-1, points: 2, grade: true, solution: false -->
```python
training_texts: List[Text] = [
    "LOL. is this u? http://supersketchyurl.com/dangerous",
    "The IRS has been trying to reach you.",
    "You won't believe what in this doc.  Click here to find out!",
    "Enclosed is your Coyote Joe's Marketplace Rewards Card."
    "Logan I'd like to add you to my professional network on LinkedIn",
]
    
t2f = TextToFeatures()
t2f.fit(training_texts)

features = t2f.transform(["Is Bill in your professional network?"]).todense()
# ensure there is one row of features for each sentence
assert features.shape[0] == 1
# ensure there are nonzero values for some selected unigram and bigram features
assert t2f.index("reach") is not None

# if a feature wasn't observed during training, it should not have an index
assert t2f.index("strawberry kangaroos") is None
assert t2f.index("Jason Mendoza") is None
```

<!-- grade_id: test-texttofeatures-2, points: 5, grade: true, solution: false -->
```python
# get the texts from the training data
examples: Iterable[Tuple[str, str]]  = read_smsspam("data/smsspam/SMSSpamCollection.train")
training_texts: List[str]            = [text for _, text in examples]

# create the feature extractor from the training texts
t2f = TextToFeatures()
t2f.fit(training_texts)

# extract features for some made-up sentences
features = t2f.transform([
    "There are some things that I need to send to you.",
    "Hello!"
]).todense()

# make sure there is one row of features for each sentence
assert len(features.shape) == 2
n_rows, n_cols = features.shape
assert n_rows == 2

# make sure there are nonzero values for some selected unigram and bigram
# features in the first sentence
# If you're getting an error here: Did you create your features such that
# the features that are checked here will actually be calculated?
indices = [t2f.index(f) for f in ["need", "to you"]]
assert len(set(indices)) > 1
row_indices, col_indices = features[:, indices].nonzero()
assert np.all(row_indices == 0)
assert len(col_indices) == 2
```

## `TextToLabels`
Using the provided skeleton, complete the class `TextToLabels` to map class labels (strings) to integers.

<!-- grade_id: code-texttolabels, points: 0, grade: false, solution: true -->
```python
class TextToLabels:
    def __init__(self):
        """
        Initializes an object for converting texts to labels.
        """
        
        # HINT: As with the previous class, you may choose to use an sklearn
        #  class here. See the documentation at
        # https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.LabelEncoder.html
        #
        # As before, be sure to add an import statement for the sklearn class
        #  that you choose in the (editable) import cell towards the top of
        #  this notebook.
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def fit(self, training_labels: Iterable[Text]) -> None:
        """
        Assigns each distinct label a unique integer.
        
        
        Training labels are analyzed to determine the vocabulary, 
        i.e., all labels that the converter will support. 
        Each such label will be associated with a unique integer index 
        that may later be accessed via the .index() method.

        :param training_labels: The training labels.
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def index(self, label: Text) -> Union[None, int]:
        """Returns the index in the vocabulary of the given label.

        :param label: A label
        :return: The unique integer index associated with the label.
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def transform(self, labels: Iterable[Text]) -> NDArray:
        """
        Creates a label vector from a sequence of labels.

        Each entry in the vector corresponds to one of the input labels. The
        value at index j is the unique integer associated with the jth label.

        :param labels: A sequence of labels.
        :return: A vector, with one entry for each label.
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        
        
    def __contains__(self, label: Text) -> bool:
        """
        Special "dunder" method to check if a label is known to the TextToLabels instance.
        
        labeler = TextToLabels()
        labeler.fit(["POSITIVE", "NEGATIVE"])

        # should be True:
        "POSITIVE" in labeler 
        
        # should be False:
        "MBOP" in labeler
        
        :return: True if the label was seen in the training data; False otherwise
        """
        # NOTE: you do not need to change this if you've implemented .index() correctly!
        return False if self.index(label) is None else True
```

## Test labels (7pts)

2 points for the first testing cell, 5 points for the second cell

<!-- grade_id: test-texttolabels-1, points: 2, grade: true, solution: false -->
```python
training_labels = ["SPAM", "NOT_SPAM"]
lbl_encoder     = TextToLabels()
lbl_encoder.fit(training_labels)

assert "SPAM" in lbl_encoder
assert lbl_encoder.index("SPAM") is not None

# this label wasn't seen in the training data labels
assert "DISSERTATION" not in lbl_encoder
assert lbl_encoder.index("DISSERTATION") is None
```

<!-- grade_id: test-texttolabels-2, points: 5, grade: true, solution: false -->
```python
def test_labels():
    # get the texts from the training data
    examples = read_smsspam("data/smsspam/SMSSpamCollection.train")
    labels = [label for label, _ in examples]

    # create the label encoder from the training texts
    lbl_encoder = TextToLabels()
    lbl_encoder.fit(labels)

    # just a simple convenience function to use for testing...
    to_labels = lambda labels: lbl_encoder.transform(labels)
    
    # make sure that some sample labels are encoded as expected
    ham_index = lbl_encoder.index("ham")
    spam_index = lbl_encoder.index("spam")
    assert ham_index != spam_index
    assert np.all(
        to_labels(["ham", "spam", "spam"]) == [ham_index, spam_index, spam_index]
    )

test_labels()
```

# `Classifier`
Using the provided skeleton, complete the class `Classifier` which will use `sklearn`'s logistic regression classifier.  You will implement the `train` and `predict` methods.

<!-- grade_id: cell-b30da2f58a08c088, points: 0, grade: false, solution: true -->
```python
class Classifier:
    def __init__(self):
        """
        Initalizes a logistic regression classifier.
        """

        # HINT: See the documentation at
        # https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html
        #  Think about the input parameter values you'll need to use.
        #
        # As before, be sure to add an import statement for the sklearn class
        #  that you choose in the (editable) import cell towards the top of
        #  this notebook.
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def train(self, features: NDArray, labels: NDArray) -> None:
        """
        Trains the classifier using the given training examples.

        :param features: A feature matrix, where each row represents a text.
        Such matrices will typically be generated via TextToFeatures.
        :param labels: A label vector, where each entry represents a label.
        Such vectors will typically be generated via TextToLabels.
        """
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # just an alias for "train"
    fit = train
    
    def predict(self, features: NDArray) -> NDArray:
        """Makes predictions for each of the given examples.

        :param features: A feature matrix, where each row represents a text.
        Such matrices will typically be generated via TextToFeatures.
        :return: A prediction vector, where each entry represents a label.
        """
        # YOUR CODE HERE
        raise NotImplementedError()
```

## Min F1 Score (4pts)

If you've gotten to this point but you're not passing these tests, it indicates that your model performance could be improved. What can you do to alter your model above, so that its performace here can be improved?

<!-- grade_id: test-predictions, points: 4, grade: true, solution: false -->
```python
def no_peeking() -> bool:
    train_data = set(read_smsspam("data/smsspam/SMSSpamCollection.train"))
    dev_data   = set(read_smsspam("data/smsspam/SMSSpamCollection.devel"))
    if len(train_data) < 10:
        print("Something is wrong with the training data")
        return False
    if len(dev_data) < 10:
        print("Something is wrong with the dev data")
        return False
    
    for dev_ex in dev_data:
        if dev_ex in train_data:
            print(dev_ex)
            print("Dev data should not overlap with train data!")
            return False
    return True
    
def test_prediction() -> Tuple[int, int]:
    # get texts and labels from the training data
    train_examples = read_smsspam("data/smsspam/SMSSpamCollection.train")
    train_labels, train_texts = zip(*train_examples)

    # get texts and labels from the development data
    dev_examples = read_smsspam("data/smsspam/SMSSpamCollection.devel")
    dev_labels, dev_texts = zip(*dev_examples)

    # fit/"train" the feature extractor and label encoder
    to_features = TextToFeatures()
    to_features.fit(train_texts)
    to_labels   = TextToLabels()
    to_labels.fit(train_labels)

    # train the classifier on the training data
    clf = Classifier()
    clf.train(to_features.transform(train_texts), to_labels.transform(train_labels))

    # make predictions on the development data
    predicted_indices = clf.predict(to_features.transform(dev_texts))
    assert np.array_equal(predicted_indices, predicted_indices.astype(bool))

    # measure performance of predictions
    dev_indices   = to_labels.transform(dev_labels)
    spam_label    = to_labels.index("spam")
    f1            = f1_score(dev_indices, predicted_indices, pos_label=spam_label)
    accuracy      = accuracy_score(dev_indices, predicted_indices)

    print(f"\n{f1:.1%} F1 and {accuracy:.1%} accuracy on SMSSpam development data")
    return f1, accuracy


# ensure the train and dev data is well-formed
assert no_peeking() is True, "Problem with train and/or dev data"

# make sure that performance is adequate
f1, accuracy = test_prediction()
min_f1, min_accuracy = 0.89, 0.97

assert accuracy > min_accuracy, f"accuracy {accuracy:.1%} did not exceed {min_accuracy:.1%}"
assert f1 > min_f1, f"f1 {f1:.1%} did not exceed {min_f1:.1%}"
```

## High F1 score (2 pts)

Adjust your classifier (ex. features, regularization, etc.) as needed to achieve a minimum F1 score of 0.94.

<!-- grade_id: test-high-performance, points: 2, grade: true, solution: false -->
```python
f1, accuracy = test_prediction()

min_f1, min_accuracy = 0.94, 0.98
# make sure that performance is adequate
assert accuracy > min_accuracy, f"accuracy {accuracy:.2%} did not exceed {min_accuracy:.2%}"
assert f1 > min_f1, f"f1 {f1:.2%} did not exceed {min_f1:.2%}"
```